# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">**PHYS 231 Lab #2**<span style="color:red"> $\to$ (5 possible marks)</span> </span></div> 

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# Ohm's Law and Kirchhoff Loop & Junction Rules

***
## Learning Objectives:
* <b><span style="color:rgba(0, 153, 76, 1);">Continue practicing setting up and using the oscilloscope to make measurements.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Verify Ohm's Law by using a Jupyter notebook to analyze a set of $V$ vs $I$ data.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Make voltage measurements by placing a voltmeter in parallel with the component of interest.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Make current measurements by placing an ammeter in series with the component of interest.</span></b>

***
## Autograding:
The PHYS 231 Jupyter Notebooks will make use of some autograding.  To make the autograding work, the two cells below need to be executed.  The first one installs any missing required packages and the second imports required packages/modules.  If 'PHYS231.Installer()' reports that some functions have been installed, the user should execute the PHYS231.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS231.Installer()' command requires the file 'PHYS231.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Clear all variables.
from IPython import get_ipython
get_ipython().run_line_magic('reset', '-sf')

# Hide unnecessary warnings.
import warnings
warnings.filterwarnings('ignore')

# Import PHYS231.py and then run the installer function.
import PHYS231
PHYS231.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 231 - Lab 2.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 231 labs.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
from IPython.display import YouTubeVideo
import math
import sys
```

```python
PHYS231.graderCheck(['PHYS231' in sys.modules, 'PHYS231' in sys.modules, 'numpy' in sys.modules], ['PHYS231', 'otter', 'numpy'], grader.check('q0.1'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Name & Student Number</span></div>  

So that we know who you were working with, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number as strings (i.e. place your names and student numbers between single or double quotes.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
PHYS231.graderCheck([last_name_1, first_name_1, student_number_1, last_name_2, first_name_2, student_number_2], ['last_name_1', 'first_name_1', 'student_number_1', 'last_name_2', 'first_name_2', 'student_number_2'], grader.check('name_and_student_number'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Introduction (5 minutes)</span></div>

## Breadboards...

In this lab, you will use a *breadboard* to build the circuits that you will study. A schematic diagram of the layout of a breadboard is shown in **Fig. 1**.  Breadboards contain arrays of terminals (holes/receptacles) that can be used to quickly assemble and prototype circuits.  

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-231-images/09080efafaa0acde3022b216a2173a2bfe9be9f2/Lab-2/images/breadboard.svg" alt="Beadboard" style="width: 20%; height: auto;" /><br>
<b>Fig. 1: Schematic diagram of a breadboad used to build and test circuits.  Image taken from FreeSVG.org: <a href = "https://freesvg.org/breadboardexplained">https://freesvg.org/breadboardexplained</a>.</b>
</center>
</p>

**Figure 1** shows that the columns of terminals are labeled $+$, $-$, a, b, c, d, e, f, g, h, i, and j.  If you look carefully at the figure, or look at the breadboard at your station, every fifth row of "terminal strips" is numbered.  Here's what you need to know about the breadboard terminals:

 - All of the terminals in a single column labelled $+$ are connected to one another.  These "power terminals" are next to a <span style="color:red">red</span> line. 
 - All of the terminals in a single column labelled $-$ are connected to one another.  These "power terminals" are next to a <span style="color:blue">blue</span> line.
 - The $+$ and $-$ power terminals are **NOT** connected to one another.  That is, they're **isolated** from one another.
 - If there are multiple columns of $+$ power terminals, they are **isolated** from one another.
 - If there are multiple columns of $-$ power terminals, they are **isolated** from one another.
 - The a, b, c, d, and e terminals in a *single* row/"terminal strip" are connected to one another.
 - The f, g, h, i, and j terminals in a *single* row/"terminal strip" are connected to one another.
 - The a, b, c, d, and e terminals in a particular row are **isolated** from the f, g, h, i, and j terminals in the same row.  Put another way, "terminal strips" on either side of a "divider" are **NOT** connected, even if they are in the same row.
 - Terminal strips in different rows are always isolated from one another.
 
Once you become comfortable with the layout of breadboards and the internal connections between terminals, they can be used to very quickly assemble and study electric circuits.

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Ohm's Law (10 minutes)</span></div>

Enter your measurements for the resistor voltage and current and their corresponding uncertainties.  You can enter the data as a list.  The syntax is as follows:
```python
ListA = [1, 2, 3, 6, 8, 12]
```

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.1</span>** **<span style="color:red">(2 marks)</span>**  

In the cell below, replace the ... with your data.  You voltage and current data should both be entered as amplitudes or root-mean-square (rms) values.

*** Please do not change anything to the left of the equals sign. ***

```python
V = ...
dV = ...
I = ...
dI = ...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.2</span>** **<span style="color:red">(1 mark)</span>**  

Next, use the 'Scatter()' function from PHYS231.py to plot $V\pm \Delta V$ as a function of $I$.

Fill in the missing data needed when calling the 'Scatter()' function.  Remember that the scatter function accepts arrays or lists as inputs.

For additional information about the 'Scatter()' function, see **Appendix A** near the end of this notebook.   

*** Please do not change anything to the left of the equals sign. ***

```python
PHYS231.Scatter(xData = ..., 
                yData = ..., 
                yErrors = ...,
                xlabel = 'current', 
                ylabel = 'voltage', 
                xUnits = 'A', 
                yUnits = 'V');
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.3</span>** **<span style="color:red">(1 mark)</span>**  

Next, we want to fit the data plotted in **<span style="color:blue">Question 2.2</span>** to a straight line. 

Replace the ... with the missing data needed when calling the 'LinearFit()' function.  Thge slope of the fit gives the value of the resistor used.  Does this value agree with what you expected?

For additional information about the 'LinearFit()' function, see **Appendix B** near the end of this notebook.   

*** Please do not change anything to the left of the equals sign. ***

```python
m, b, dm, db, fig = PHYS231.LinearFit(xData = ..., 
                                      yData = ...,  
                                      yErrors = ...,
                                      xlabel = 'current', 
                                           ylabel = 'voltage', 
                                           xUnits = 'A', 
                                           yUnits = 'V')
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Submission (5 minutes)</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.1:</span>**  

We welcome your feedback on the PHYS 231 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 121 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 3.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file using the link that appears **below** the ```grader.check_all(...)``` cell.
- $\Leftarrow$ Please do **NOT** download the zip file from the frame along the left-hand side of the screen. 
- Upload the .zip file to the PHYS 121 Canvas gradebook.
- **Do NOT change the name of the .zip file.**
- **Do NOT modify the contents of the .zip file.**

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
display(grader.check_all())
PHYS231.saveMessage()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
PHYS231.waitMessage()
grader.export(files=["PHYS121_DataLogger.txt"])
PHYS231.proceedMessage()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part  - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.  
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.  
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS231
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS231.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix B &ndash; The LinearFit Function...</div>

The linear-fit function is called as follows:
```python
PHYS231.LinearFit(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False)
```
The xData and yData lists are required.  The others are optional with default values set.  The function returns the following outputs:
```python
Slope, Yintercept, errSlope, errYintercept, fig
```
The first four outputs are numerical values and 'fig' is the formatted plot. 

If $y$-uncertainities are provided, the function will perform a weighted fit.  The $y$-uncertainties list must be the same length as $x$- and $y$-data lists.  If $y$-uncertainties are not provided, the fit will be unweighted.  The other optional arguments include $x$- and $y$-axis names and units for the $x$- and $y$-datasets.  These must be entered as strings (enclosed in quotations) and they are used for formatting the outputs of the function.  If 'fill' is set to True, shading will be added around the best-fit line representing confindence intervals. 

### Unweighted Fit Example Implmentation
The code block below shows the most basic use of 'LinearFit' for an unweighted fit without any of the optional arguments.
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
m, b, dm, db, fig = PHYS231.LinearFit(V, I)
```

### Weighted Fit Example Implmentation
This second blcok of code shows how to use 'LinearFit' to do a weighted fit.  It also makes use of all of the other optional arguments.  
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
errI = [0.005, 0.012, 0.020, 0.025]
m, b, dm, db, fig = PHYS231.LinearFit(V, I, errI, 'voltage' , 'current', 'V', 'A', True)
```

The 'm, b, dm, db, fig = ...' syntax is used to store the function outputs in varibles.  If you're interested in details of the fitting process, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Weighted%20Linear%20Fit.html
***

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: September 14, 2024

Tests for name_and_student_number (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...


Tests for q0.1 (points: None)
```
>>> 'PHYS231' in sys.modules and 'otter' in sys.modules and ('numpy' in sys.modules)
True
```
failure_message: Make sure you execute the cells containing "import PHYS231" and "import otter" as well as the cell beneath the heading "Import Modules".
